
<problem>
  <customresponse cfn="check_function">
    <script type="loncapa/python">
      <![CDATA[ template-file(check_function.py) ]]>
    </script>
    
    <p>This is paragraph text displayed before the iframe.</p>
    
    <jsinput
        gradefn="JSInputDemo.getGrade"
        get_statefn="JSInputDemo.getState"
        set_statefn="JSInputDemo.setState"
        initial_state='{"selectedChoice": "incorrect1", 
                       "availableChoices": ["incorrect1", "correct", "incorrect2"]}'
        width="600"
        height="100"
        html_file="https://files.edx.org/custom-js-example/jsinput_example.html"
        title="Dropdown with Dynamic Text"
        sop="false"
        />
  </customresponse>
</problem>



<!-- <!DOCTYPE html> -->
<!-- <html lang="en"> -->
<!--   <head> -->
<!--     <title>Dropdown with Dynamic Text</title> -->
<!--     <link rel="stylesheet" type="text/css" href="https://files.edx.org/custom-js-example/jsinput_example.css"> -->
<!--   </head> -->
<!--   <body> -->
<!--     <script src="https://files.edx.org/custom-js-example/jschannel.js"></script> -->
<!--     <script src="https://files.edx.org/custom-js-example/jsinput_example.js" defer></script> -->
<!--     <label class="directions">Select an option from the list: -->
<!--         <select class="choices"></select> -->
<!--     </label> -->
<!--     <p aria-live="polite" class="feedback"></p> -->
<!--   </body> -->
<!-- </html> -->
